# Introducción a Jupyter

## Expresiones aritmeticas y algebraicas

Empezaremos esta práctica con algo de conocimientos previos de programación. Se que muchos de ustedes no han tenido la oportunidad de utilizar Python como lenguaje de programación y mucho menos Jupyter como ambiente de desarrollo para computo cientifico, asi que el primer objetivo de esta práctica será acostumbrarnos a la sintaxis del lenguaje y a las funciones que hacen especial a Jupyter.

Primero tratemos de evaluar una expresión aritmetica. Para correr el código en la siguiente celda, tan solo tienes que hacer clic en cualquier punto de ella y presionar las teclas Shift + Return.

```python
2 + 3
```

```python
2*3
```

```python
2**3
```

```python
sin(pi)
```

Sin embargo no existen funciones trigonométricas cargadas por default. Para esto tenemos que importarlas de la libreria ```math```:

```python
from math import sin, pi
sin(pi)
```

## Variables

Las variables pueden ser utilizadas en cualquier momento, sin necesidad de declararlas, tan solo usalas!

```python
a = 10
```

```python
a
```

### Ejercicio

Ejecuta el siguiente calculo y guardalo en una variable ```c```:

$$
c = \pi *10^2
$$

> Nota: Una vez que hayas concluido el calculo y guardado el valor en una variable, despliega el valor de la variable al ejecutar en una celda el nombre de la variable

<!-- grade_id: cell-78c2e2820d00f49c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION HERE
from numpy import pi
c = pi*10**2
### END SOLUTION HERE
```

<!-- grade_id: cell-f8da2ce5baac53e0, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION HERE
c
### END SOLUTION HERE
```

Ejecuta la prueba de abajo para saber si has creado el codigo correcto

<!-- grade_id: cell-87bee619e6dfd6b7, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(_, c)
print("Sin errores")
```

## Listas

Las listas son una manera de guardar varios datos en un mismo arreglo. Podemos tener por ejemplo:

```python
A = [2, 4, 8, 10]
A
```

Pero si intentamos multiplicar estos datos por un numero, no tendrá el comportamiento esperado.

```python
A*2
```

## Funciones

Podemos definir funciones propias de la siguiente manera:

```python
f = lambda x: x**2 + 1
```

Esta linea de codigo es equivalente a definir una función matemática de la siguiente manera:

$$
f(x) = x^2 + 1
$$

Por lo que si la evaluamos con $x = 2$, obviamente obtendremos como resultado $5$.

```python
f(2)
```

Esta notación que introducimos es muy util para funciones matemáticas, pero esto nos obliga a pensar en las definiciones de una manera funcional, lo cual no siempre es la solución (sobre todo en un lenguaje con un paradigma de programación orientado a objetos).

Esta función tambien puede ser escrita de la siguiente manera:

```python
def g(x):
    y = x**2 + 1
    return y
```

Con los mismos resultados:

```python
g(2)
```

### Ejercicio

Define una función que convierta grados Celsius a grados Farenheit, de acuerdo a la siguiente formula:

$$
F = \frac{9}{5} C + 32
$$

<!-- grade_id: cell-28a81718d17ce3d2, points: 0, grade: false, solution: true -->
```python
def cel_a_faren(grados_cel):
    ### BEGIN SOLUTION HERE
    grados_faren = (9/5)*grados_cel + 32
    ### END SOLUTION HERE
    return grados_faren
```

Y para probar trata de convertir algunos datos:

```python
cel_a_faren(-1)
```

<!-- grade_id: cell-14df3266209f3b02, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(cel_a_faren(10), 50)
assert_equal(cel_a_faren(50), 122)
### BEGIN HIDDEN TESTS
assert_equal(cel_a_faren(0), 32)
assert_equal(cel_a_faren(-1), 30.2)
### END HIDDEN TESTS
print("Sin errores")
```

## Ciclos de control

Cuando queremos ejecutar código varias veces tenemos varias opciones, vamos a explorar rapidamente el ciclo for.
```python
    for paso in pasos:
        ...
        codigo_a_ejecutar(paso)
        ...
```
En este caso el codigo se ejecutará tantas veces sean necesarias para usar todos los elementos que hay en pasos.

Por ejemplo, pordemos ejecutar la multiplicacion por 2 en cada uno de los datos:

```python
for dato in A:
    print(dato*2)
```

ó agregarlo en una lista nueva:

```python
B = []
for dato in A:
    B.append(dato*2)
    
B
```

### Ejercicio

* Crea una lista ```C``` con los enteros positivos de un solo digito, es decir: $\left\{ x \in \mathbb{Z} \mid 0 \leq x < 10\right\}$
* Crea una segunda lista ```D``` con los cuadrados de cada elemento de ```C```

```python
C = [] # Escribe el codigo para declarar el primer arreglo adentro de los corchetes

C
```

```python
D = []

# Escribe el codigo de tu ciclo for aqui

D
```

Ejecuta las pruebas de abajo

```python
from pruebas_1 import prueba_1_3
prueba_1_3(C, D)
```

## Matrices

Para trabajar con matrices el flujo de trabajo es un poco diferente, ya que los arreglos manejan un conjunto diferente de operaciones; afortunadamente la librería `numpy` ya tiene definido un objeto de tipo matriz, con el cual podemos hacer las operaciones que queramos:

```python
from numpy import matrix
```

```python
A = matrix([[1, 2], [3, 4]])
A
```

```python
v1 = matrix([[1], [2]])
v1
```

```python
# Dependiendo de la version de python que exista en tu computadora,
# esta operacion pudiera no funcionar, en dado caso solo hay que 
# cambiar @ por *
A@v1
```

```python
# La siguiente linea no va a funcionar, porque?
v1@A
```

# Ejercicio

* Declara una matriz de rotación ```rot```, en dos dimensiones con un angulo de 30º
* Aplica esta matriz de rotación al vector ```vec``` definido como $\begin{pmatrix} 2 \\ 2\end{pmatrix}$

<!-- grade_id: cell-4ba6009db4e4d822, points: 0, grade: false, solution: true -->
```python
from numpy import sin, cos, pi
τ = 2*pi
### BEGIN SOLUTION HERE
from numpy import matrix
rot = matrix([[cos(τ/12), -sin(τ/12)], [sin(τ/12), cos(τ/12)]])
vec = matrix([[2], [2]])

vec_rot = rot@vec
### END SOLUTION HERE
vec_rot
```

<!-- grade_id: cell-5cec95c955823778, points: 2, grade: true, solution: false -->
```python
from numpy.testing import assert_array_equal
assert_array_equal(vec_rot, matrix([[2*(cos(τ/12)-sin(τ/12))], [2*(cos(τ/12)+sin(τ/12))]]))
### BEGIN HIDDEN TESTS
assert_array_equal(rot, matrix([[cos(τ/12), -sin(τ/12)], [sin(τ/12), cos(τ/12)]]))
assert_array_equal(vec, matrix([[2], [2]]))
### END HIDDEN TESTS
print("Sin errores")
```