# Problemas

Crea una función que sea capaz de calcular la matriz de rotación en un plano con cualquier angulo que se le de.

<!-- grade_id: cell-cbe6b534af68b731, points: 0, grade: false, solution: true -->
```python
def rotacion(θ):
    ### BEGIN SOLUTION HERE
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), -sin(θ)],
                [sin(θ), cos(θ)]])
    return A
    ### END SOLUTION HERE
```

<!-- grade_id: cell-3767258f4db444e1, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import matrix
assert_allclose(rotacion(0), matrix([[1, 0], [0, 1]]), rtol=1e-5, atol=1e-5)
### BEGIN HIDDEN TESTS
from numpy import pi, sin, cos
assert_allclose(rotacion(pi), matrix([[-1, 0], [0, -1]]), rtol=1e-5, atol=1e-5)
assert_allclose(rotacion(pi/6), matrix([[cos(pi/6), -sin(pi/6)], [sin(pi/6), cos(pi/6)]]), rtol=1e-5, atol=1e-5)
### END HIDDEN TESTS
```

Calcule la matriz de rotación en un plano, con un angulo de 30º, 50º y 100º, y guardalos en las variables ```A1```, ```A2```, ```A3```.

<!-- grade_id: cell-5856c502c0bdac7c, points: 0, grade: false, solution: true -->
```python
from numpy import pi
τ = 2*pi
### BEGIN SOLUTION HERE
A1 = rotacion(τ*30/360)
A2 = rotacion(τ*50/360)
A3 = rotacion(τ*100/360)
### END SOLUTION HERE
```

<!-- grade_id: cell-c84ab72c9098ebac, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
### BEGIN HIDDEN TESTS
from numpy import pi
τ = 2*pi
def prueba_rotacion(θ):
    from numpy import matrix, sin, cos
    A = matrix([[cos(θ), -sin(θ)],
                [sin(θ), cos(θ)]])
    return A
assert_allclose(A1, prueba_rotacion(τ*30/360), rtol=1e-5, atol=1e-5)
assert_allclose(A2, prueba_rotacion(τ*50/360), rtol=1e-5, atol=1e-5)
assert_allclose(A3, prueba_rotacion(τ*100/360), rtol=1e-5, atol=1e-5)
### END HIDDEN TESTS
```

Aplique las matrices de rotación de 30º, 100º y 50º al vector $\begin{pmatrix} 2 \\ 2\end{pmatrix}$ y guarde el resultado en la variable ```res1```.

<!-- grade_id: cell-44ddf2f4d6904fa6, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION HERE
vec = matrix([[2], [2]])
res1 = A2@A3@A1@vec
### END SOLUTION HERE
```

<!-- grade_id: cell-54cd99b9320ea064, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
### BEGIN HIDDEN TESTS
vec_prueba = matrix([[2], [2]])
assert_allclose(res1, A2@A3@A1@vec_prueba, rtol=1e-5, atol=1e-5)
### END HIDDEN TESTS
```

Aplique las matrices de rotación de 50º, 100º y 30º al vector $\begin{pmatrix} 2 \\ 2\end{pmatrix}$ y guarde el resultado en la variable ```res2```.

<!-- grade_id: cell-1f99db1538761d2f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION HERE
vec = matrix([[2], [2]])
res1 = A1@A3@A2@vec
### END SOLUTION HERE
```

<!-- grade_id: cell-7423316a8b222d0c, points: 1, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
### BEGIN HIDDEN TESTS
vec_prueba = matrix([[2], [2]])
assert_allclose(res1, A1@A3@A2@vec_prueba, rtol=1e-5, atol=1e-5)
### END HIDDEN TESTS
```